<problem display_name="Question 4 Response" markdown="null">
  <text>
    <p>
      <h4>Problem With Adaptive Hint</h4>
    </p>
    <p>
Enter the average number of cookies sold in classrooms A and B </p>
    <script type="text/python" system_path="python_lib">
def test_str(expect, ans):
  print expect, ans
  ans = ans.strip("'")
  ans = ans.strip('"')
  return expect == ans.lower()

def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid]).lower()
  print 'hint_fn called, ans=', ans
  hint = ''
  if hint:
    hint = "&lt;font color='blue'&gt;Hint: {0}&lt;/font&gt;".format(hint)
    new_cmap.set_hint_and_mode(aid,hint,'always')
</script>
    <label>
   <p style="display:inline"><i>School average</i> = 
    <customresponse cfn="test_str" expect="python">
        <textline correct_answer="python"/>
        <hintgroup hintfn="hint_fn"/>
      </customresponse>
    <p style="display:inline"/>
  </p>
      
    </label>
  </text>
</problem>
